```scala
import $file.hw2stdlib_new
import hw2stdlib_new._
```

Output:
```
[32mimport [39m[36m$file.$            
[39m
[32mimport [39m[36mhw2stdlib_new._[39m
```

# Homework 2

Due 9/12 at 11:59pm

## Submission Instructions
Upload only this .ipynb file to Canvas. Do not add anything to hw2stdlib since you can't submit it.

In this homework we will develop more useful functions for numbers as well as some of the functions on lists that we spoke about during lecture.

## Arithmetic Functions

### Problem 1 (5 Points)

Write the minus function, which subtracts a natural number from another natural number. If the result would be less than zero then leave the answer as zero.


<!-- grade_id: cell-19061b6ebc83ee13, points: 0, grade: false, solution: true -->
```scala
def minus(x : Nat, y : Nat) : Nat = {
    // matching on x and y
    (x, y) match {
        // case x == 0 and y >= 0, we return 0 because we don't have negative numbers ---> base case 
        case (Zero, _)          => Zero
        // case y == 0 and x > 0 ( x == Succ(x) ), we return x ( Succ(x) )            ---> base case
        case (Succ(x),Zero)     => Succ(x)
        // case x > 0 ( x == Succ(x) ) and y > 0 ( y == Succ(y) ), we return minus(y-1, x-1)
        case (Succ(x), Succ(y)) => minus(x, y)
    }
}
```

Output:
```
defined [32mfunction[39m [36mminus[39m
```

```scala
// personal tests
print_nat(minus(four, two))
print_nat(minus(four, Zero))
print_nat(minus(Zero, four))
print_nat(minus(Zero, Zero))
print_nat(minus(five, six))
print_nat(minus(two, two))
```

Output:
```
[36mres20_0[39m: [32mString[39m = [32m"2"[39m
[36mres20_1[39m: [32mString[39m = [32m"4"[39m
[36mres20_2[39m: [32mString[39m = [32m"0"[39m
[36mres20_3[39m: [32mString[39m = [32m"0"[39m
[36mres20_4[39m: [32mString[39m = [32m"0"[39m
[36mres20_5[39m: [32mString[39m = [32m"0"[39m
```

<!-- grade_id: cell-80b015bfaabd7c53, points: 3, grade: true, solution: false -->
```scala
assert(minus(four, two) == two, 1)
assert(minus(four, Zero) == four, 2)
assert(minus(Zero, four) == Zero, 3)
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
```

<!-- grade_id: cell-3fb81b3d6c3bdcd6, points: 2, grade: true, solution: false -->
```scala
// HIDDEN TEST (2 pts)
```

### Problem 2 (5 points)

Define the $\texttt{pow}$ function, which raises a natural number to some power. Notationally: 

$$
f\ x\ y = x^y
$$

Hint: Use the mult function from the stdlib for help

<!-- grade_id: cell-cf03509753bbf640, points: 0, grade: false, solution: true -->
```scala
def pow(x: Nat, y: Nat) : Nat = {
    // matching on x and y
    (x, y) match {
            // case x and y are both greater than zero, multiply x by pow(x, y-1)
            case (Succ(x), Succ(y)) => mult(Succ(x), pow(Succ(x), y))
            // case x is greater than 0 and y is 0, we return 1
            case (Succ(x), Zero)    => one
            // case x is greater than 0 and y is 1, we finally return the multiplication of x and y
            case (Succ(x), one)     => mult(Succ(x), y)
        }
}
```

Output:
```
defined [32mfunction[39m [36mpow[39m
```

```scala
// personal tests
print_nat(pow(three, two))
print_nat(pow(four, Zero))
print_nat(pow(two, two))
print_nat(pow(five, one))
```

Output:
```
[36mres31_0[39m: [32mString[39m = [32m"9"[39m
[36mres31_1[39m: [32mString[39m = [32m"1"[39m
[36mres31_2[39m: [32mString[39m = [32m"4"[39m
[36mres31_3[39m: [32mString[39m = [32m"5"[39m
```

<!-- grade_id: cell-cc0a85f98f9550b1, points: 4, grade: true, solution: false -->
```scala
assert(pow(three, two) == nine, 1)
assert(pow(four, Zero) == one, 2)
assert(pow(two, two) == four, 3)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-e4f7652875d80790, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pt)
```

### Problem 3 (5 points)

Write the "equals to" function for numbers. This function should return $True$ if $x == y$ and $False$ otherwise.

<!-- grade_id: cell-5c44d445fb4cbec0, points: 0, grade: false, solution: true -->
```scala
def eq_nat(x : Nat, y : Nat) : Bool = {
    // matching on x and y
    (x, y) match {
        // case Zero and Zero, we return True
        case (Zero, Zero)       => True
        // case both numbers are greater than zero still, we try eq_nat with x-1 and y-1
        case (Succ(x), Succ(y)) => eq_nat(x, y)
        // any other case is False
        case (_, _)             => False
    }
}
```

Output:
```
defined [32mfunction[39m [36meq_nat[39m
```

```scala
// personal tests
print(eq_nat(Zero, Zero))
print("\n")
print(eq_nat(four, Zero))
print("\n")
print(eq_nat(two, two))
print("\n")
print(eq_nat(Zero, one))
```

Output:
```
True
False
True
False
```

<!-- grade_id: cell-cc0a85f98f9550b2, points: 4, grade: true, solution: false -->
```scala
assert(eq_nat(four, four) == True, 1)
assert(eq_nat(four, Zero) == False, 2)
assert(eq_nat(eight, nine) == False, 3)
assert(eq_nat(nine, nine) == True, 4)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-e4f7652875d80791, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pt)
```

## Functions on Number Lists

### Problem 4 (5 pts)

Write the sum function. This function should take a list of natural numbers and return their sum. Here is the type:

$$
sum : \text{List}\ \mathbb{N} \rightarrow \mathbb{N}
$$

Hint: An empty list should have sum 0.

<!-- grade_id: cell-1026d25f142fc212, points: 0, grade: false, solution: true -->
```scala
def sum(xs: List[Nat]) : Nat = {
    // matching on xs
    xs match {
        // Hint: An empty list should have sum 0
        case Empty => Zero
        // else xs has elements so add x the head to the sum of xs the list without its head x
        case Cons(x, xs) => plus(x, sum(xs))
    }
}
```

Output:
```
defined [32mfunction[39m [36msum[39m
```

```scala
// personal tests
val list1: List[Nat] = Cons(one, Empty)
val list2: List[Nat] = Cons(one, Cons(seven, Cons(two, Empty)))
val list3: List[Nat] = Cons(five, list2)

// print_nat(sum(list1))
// print_nat(sum(list2))
// print_nat(sum(list3))
```

Output:
```
[36mlist1[39m: [32mList[39m[[32mNat[39m] = [33mCons[39m([33mSucc[39m(Zero), Empty)
[36mlist2[39m: [32mList[39m[[32mNat[39m] = [33mCons[39m(
  [33mSucc[39m(Zero),
  [33mCons[39m(
    [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))))),
    [33mCons[39m([33mSucc[39m([33mSucc[39m(Zero)), Empty)
  )
)
[36mlist3[39m: [32mList[39m[[32mNat[39m] = [33mCons[39m(
  [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))),
  [33mCons[39m(
    [33mSucc[39m(Zero),
    [33mCons[39m(
      [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))))),
      [33mCons[39m([33mSucc[39m([33mSucc[39m(Zero)), Empty)
    )
  )
)
```

<!-- grade_id: cell-c8396ec50b153d54, points: 4, grade: true, solution: false -->
```scala
assert(sum(Cons(one, Empty)) == one, 1)
assert(sum(Cons(one, Cons(seven, Cons(two, Empty)))) == ten, 2)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-d118cae84cab07ee, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pt)
```

### Problem 5 (5 pts)

Write the product function. This function should take alist of natural numbers and return their product. Here is the type:

$$
prod : \text{List}\ \mathbb{N} \rightarrow \mathbb{N}
$$

Hint: An empty list should have product 1.

<!-- grade_id: cell-6a2caa631bc53781, points: 0, grade: false, solution: true -->
```scala
def prod(xs: List[Nat]) : Nat = {
    // matching on xs
    xs match {
        // Hint: An empty list should have product xs
        case Empty       => one
        // else xs has elements, so we multiply x with the prod of xs which is xs without its head x
        case Cons(x, xs) => mult(x, prod(xs))
    }
}
```

Output:
```
defined [32mfunction[39m [36mprod[39m
```

<!-- grade_id: cell-cc0a85f98f9550b3, points: 4, grade: true, solution: false -->
```scala
assert(prod(Cons(one, Empty)) == one, 1)
val fourteen = Succ(Succ(Succ(Succ(ten))))
assert(prod(Cons(one, Cons(seven, Cons(two, Empty)))) == fourteen, 2)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36mfourteen[39m: [32mSucc[39m = [33mSucc[39m(
  [33mSucc[39m(
    [33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m([33mSucc[39m(Zero))))))))))))
  )
)
```

<!-- grade_id: cell-e4f7652875d80792, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pt)
```

### Problem 6 (6 pts)

Write the function that tells us if a given natural number is in a list of natural numbers. (you will need to use the equals function defined in problem 3)

$$
in : \mathbb{N} \rightarrow \text{List}\ \mathbb{N} \rightarrow \mathbb{B}
$$

<!-- grade_id: cell-603f77bab9b8f27a, points: 0, grade: false, solution: true -->
```scala
def in(n: Nat, xs: List[Nat]) : Bool = {
    // matching on xs
    l match {
        // if its empty, n is trivially not in xs
        case Empty       => False
        // else xs has elements, so we check if the head is equal to n and match on the output of this function
        case Cons(x, xs) => eq_nat(x, n) match {
            // if x == n, we can return True
            case True  => True
            // else we have to keep ckecking, so we call in on the xs without the head x
            case False => in(n, xs)
        }
    }
}
```

Output:
```
defined [32mfunction[39m [36min[39m
```

<!-- grade_id: cell-b9d72e55abc06dd4, points: 6, grade: true, solution: false -->
```scala
assert(in(one, Cons(one, Empty)) == True, 1)
assert(in(seven, Cons(one, Cons(seven, Cons(two, Empty)))) == True, 2)
assert(in(eight, Cons(one, Cons(seven, Cons(two, Empty)))) == False, 3)
assert(in(seven, Empty) == False, 4)
passed(6)
```

Output:
```
*** Tests Passed (6 points) ***
```

```scala

```